Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

---
# Lab 4: If and While
---
**This is a pair-programming lab.**

**You should do all portions of this lab with your partner, including design of the algorithm and implementation. You should not divide the problems in this lab and work on them individually; all of the work should be done with both people present and actively engaged.**

The purpose of this lab is to introduce you to while loops and expose you to pair programming for the first time in this class. This lab will also get you used to writing code that is a little longer than code from other labs. For these exercises, you and your partner should work together to create a design before you get started coding. For exercises 5, 6, and 7, the autograder cannot assist in checking your work, because the code you are writing is interactive. You will have to check these manually.

**This lab represents an increase in complexity. If you're feeling stuck, don't hesitate to reach out for help.**

---
This lab introduces:
    
    while loops
    
    random module
    
And provides practice for:

    if statements
    
    design 
    
    sums and counters in loops



---
### Exercise 1
---
Write a function called **valid_date** that takes 3 parameters: A month, day, 
and year, in that order. This function should return **True** if it is a valid
date and **False** if it is **not** a valid date.

You will need to check that the month is between 1 and 12. For different months, there are different ranges of days that are valid, be sure to account for this. Additionally, the valid ranges of days for the month February will vary depending on what year it is. Use the function **is_leap_year** from Lab 3 to determine this. (You can copy that function into the cell. If you weren't able to get isLeapYear working last time, you can try again here.)

<!-- grade_id: Part1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
### Exercise 2
---
As a continuation of the exercise above, write a function called **get_day_number**. It should take the same parameters as **valid_date** and return what day of the year the date is. (1 to 365 on a regular year and 1 to 366 on a leap year)

To calculate the day number (assume using the variable **day_num**), you could use the following three steps:
    1. day_num = 31 * (month-1) + day
    2. if the month is after February, subtract (4 * month + 23) // 10 from day_num
    3. if it's a leap year and after February 29th, add 1 to day_num

Be sure call **valid_date** inside **get_day_number** to make sure it is a valid date. If the date is invalid, return **False**.

<!-- grade_id: Part2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
### Exercise 3
---
**If n is a positive whole number greater than 2, we say that n is *prime* if no number between 2 and the square root of n (inclusive) evenly divides n.** 
For example, 5 is prime because no number aside from 1 and 5 goes into 5 evenly. 8 is not prime, because 2 goes into 8 evenly. 57 is not prime, because 3 goes into it evenly, but 59 is prime.

In the cell below, write a function called **is_prime** that takes a value of n as a parameter and returns True if n is prime and False otherwise. Use a while loop in your solution. Your solution should not loop more times than is necessary to confirm that the number is or is not prime.

To find the square root of the parameter, you can use the math.sqrt() function. If you

    import math
    
At the top of the cell, you will have access to this function and can use it to find the square root of a number.

<!-- grade_id: Part3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
### Exercise 4
---
There is a geometric series: 1 + 1/2 + 1/4 + 1/8 + 1/16 + 1/32 + ... + 1/(2^n). Every time the next, fraction is added, the sum gets closer and closer to 2. 

Write a function called **geometric_series** that takes one parameter that must be between 0 and 2. The function should calculate and return how many additions it takes to reach the parameter, including the first 1. One way to do this is to use a while loop and to increment a counter variable every time the while loop runs.

**Example**

If the parameter is 1.80, the returned value shoud be 4.

Here is why:

    1.) The first number to be added to the sum is 1, total sum is 1

    2.) The second number to be added to the sum is 1/2(0.5), total sum is 1.5

    3.) The third number to be added to the sum is 1/4(0.25), total sum is 1.75

    4.) The fourth number to be added to the sum is 1/8(0.125), total  sum is 1.875.

It takes four steps for the sum the reach 1.8

Think carefully about what values you will need to keep track of as the while loop runs.

**Note:**

Here is a short Khan Academy video about this geometric series: https://www.khanacademy.org/math/algebra2/sequences-and-series/copy-of-geometric-sequence-series/v/series-as-sum-of-sequence

**Note 2:**

You will be using this in a future exercise. 

<!-- grade_id: Part4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
### Exercise 5
---
Write a function called **guessing_game** in which the user can play a guessing game against the computer. The function will take one parameter, a maximum value. This game will prompt the user for an integer between 0 and that maximum value (inclusive) as a guess. The correct answer should be determined randomly. For example, if guessing_game(20) is called, then the maximum value that can be generated should be 20.

Use the **randrange** or **randint** function from the **random** pythom library in order to generate the correct guess. 
https://docs.python.org/3/library/random.html#random.randrange

Be sure to print whether or not the user guessed correctly. If they didn't guess correctly, communicate to the user what the correct guess would have been. (This game is not very fun, but we will design more interesting versions later.)

**Note:** Prompting the user means the function will take user input rather than a parameter. You will use the `input()` function. Since user input is taken as a string, you will need to change it into an integer using the `int()` function. This function and the ones like it cannot be automatically validated; you will need to test it yourself.

<!-- grade_id: Part5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
### Exercise 6
---
Write a function called **slightly_better_guessing_game**. It should work the same as **guessing_game**, but it will continue to give the user guesses until the user has guessed the correct number. 

Be sure to communicate with the user when they have won and when they have guessed incorrectly and need to guess again. 

**Note**: Use a while loop so that the game is still played while the user has yet to guess the correct number. 

<!-- grade_id: Part6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

---
### Exercise 7
---
Write a function called **rock_paper_scissors** in which a user can play the game rock paper scissors against a computer. Your function should take no parameters.

Take user input using the statement: "rock, paper, scissors, shoot!". Design the game such that the game will **not** end until the user **beats** the computer (use a while loop for this).

The easiest way to do this might be to use the random.choice function. You can also use the randrange or randint function from the **random** pythom library in order to generate the computer's move. 
https://docs.python.org/3/library/random.html#random.randrange

Print **what the computer played**, and print **whether the user lost, tied or won** the game. 

<!-- grade_id: Part7, points: 0, grade: false, solution: true -->
```python
import random
def  rock_paper_scissors():
    game = True
    while game:
        user_choice = input ("rock, paper, scissors, shoot!")
        computer = ["rock","paper","scissors"]
        r = random.choice (computer)
        print ("what the computer played:", r)
        if user_choice == "paper" and r == "scissors": 
            print ("you lose")
            game = True
        if user_choice == "rock" and r == "paper":
            print ("You lose :(")
            game = True
        if user_choice == "scissors" and r == "rock":
            print ("You lose :(")
            game = True
        if user_choice == "rock" and r == "scissors":
            print ("You Win!")
            game = False
        if user_choice == "scissors" and r == "paper":
            print ("You Win!")
            game = False
        if user_choice == "paper" and r == "rock":
            print ("You Win!")
            game = False
        if user_choice == r:
            print ("It's a Tie!")
            game = False
        
rock_paper_scissors()
    
```

<!-- grade_id: Test1-case0, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
from nose.tools import assert_equal
### BEGIN TEST
assert_equal(is_leap_year(2000), True)
assert_equal(is_leap_year(2016), True)
assert_equal(is_leap_year(1900), False)
assert_equal(is_leap_year(2015), False)
### END TEST
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-9-9c1013bb70f5> in <module>
      3 from nose.tools import assert_equal
      4 ### BEGIN TEST
----> 5 assert_equal(is_leap_year(2000), True)
      6 assert_equal(is_leap_year(2016), True)
      7 assert_equal(is_leap_year(1900), False)

NameError: name 'is_leap_year' is not defined
```

<!-- grade_id: Test1-case1, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(13, 29, 1200), False)
### END TEST
```

<!-- grade_id: Test1-case2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(10, 20, 0), False)
### END TEST
```

<!-- grade_id: Test1-case3, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(1, 31, 2004), True)
### END TEST
```

<!-- grade_id: Test1-case4, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(5, 32, 2004), False)
### END TEST
```

<!-- grade_id: Test1-case5, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(2, 29, 1200), True)
### END TEST
```

<!-- grade_id: Test1-case6, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(2, 30, 1200), False)
### END TEST
```

<!-- grade_id: Test1-case7, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(2, 28, 1201), True)
### END TEST
```

<!-- grade_id: Test1-case8, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(2, 29, 1201), False)
### END TEST
```

<!-- grade_id: Test1-case9, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(11, 31, 1901), False)
### END TEST
```

<!-- grade_id: Test1-case10, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(11, 30, 1901), True)
### END TEST
```

<!-- grade_id: Test1-case11, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(valid_date(-1, -30, 1901), False)
### END TEST
```

<!-- grade_id: Test2-case1, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(get_day_number(2, 1, 2016), 32)
### END TEST
```

<!-- grade_id: Test2-case2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(get_day_number(12, 31, 2015), 365)
### END TEST
```

<!-- grade_id: Test2-case3, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(get_day_number(12, 31, 2016), 366)
### END TEST
```

<!-- grade_id: Test2-case4, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(get_day_number(13, 15, 2016), False)
### END TEST
```

<!-- grade_id: Test2-case5, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(get_day_number(12, -1, 2016), False)
### END TEST
```

<!-- grade_id: Test3-Case1, points: 3, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(is_prime(101), True)
assert_equal(is_prime(103), True)
assert_equal(is_prime(17), True)
### END TEST
```

<!-- grade_id: Test3-Case2, points: 3, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(is_prime(105), False)
assert_equal(is_prime(51), False)
### END TEST
```

<!-- grade_id: Test4-Case1, points: 3, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(geometric_series(1.89), 5)
### END TEST
```

<!-- grade_id: Test4-Case2, points: 3, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(geometric_series(1.999999), 21)
### END TEST
```

<!-- grade_id: Test5, points: 6, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
# Test for guessing_game
### BEGIN TEST
### END TEST
```

<!-- grade_id: Test6, points: 6, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
# Test for slightly_better_guessing_game
### BEGIN TEST
### END TEST
```

<!-- grade_id: Test7, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
# Test for rock, paper, scissors
### BEGIN TESTS
#rock_paper_scissors()
### END TEST
```

```python

```