Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ihsan Alaeddin"
COLLABORATORS = "Group 1"
```

---

---
# Lab 7: Files
---

In this lab, you will be introduced to working with files. This entails:
    
* opening files
    
* reading files
    
* manipulating the contents of the file
    
* writing files
    
You will also get more practice with:
    
* working with strings
    
* if, elif, else
    
* command line
    
* loops
    
---

REMEMBER:

1. After you open a file in your code, ***close*** it in your code after you are done with it (or use the "with" structure)! Otherwise, you may have problems with opening it in the future (directly in Jupyter or through Python). 

2. You can't open a file unless it is in the directory of this lab, so you need to go through the terminal and make sure whatever files you're using are in the same directory as the lab.

It may be useful to view the contents of a file before you get started on a design. To do so, you can use the `cat` command in the terminal.

Here is an outside resource that may come in handy:

http://www.pythonforbeginners.com/files/reading-and-writing-files-in-python

---
### Exercise 1
---
Expand your code from Exercise 8 from Lab 6 on Strings (namely, the is_palindrome function) so that it takes the name of a text file as a parameter rather than a string. If the contents of the file form a palindrome, the function should return **True**; otherwise, it should return **False**. Remember that only alphabetical letters are considered when checking if a string is a palindrome or not.

(You should copy the code of is_palindrome from your Lab 6 to this assignment and modify it here. Do not change the previous assignment. You may need to copy over other functions as well, if they are used by is_palindrome. If your implementation of is_palindrome did not pass all test cases in Lab 6, then you should fix all these problems before you continue working on this exercise.)

The name of the function should be the same as before, **is_palindrome**. There are two files with the following paths. One is a palindrome and one isn't. You should copy them to this assignment's directory using the ***terminal***. Once they are there, you can use them for testing.

    ~barbeda/cs128-share/notPalindrome.txt
    
    ~barbeda/cs128-share/224wordPalindrome.txt
    
Once your files are copied over, you should be able to call your new function with 

    is_palindrome("notPalindrome.txt")
    is_palindrome("224wordPalindrome.txt")

<!-- grade_id: Part1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def is_palindrome (e):
    file_1 = open(e,"r")
    we = file_1.read()
    for line in we:
        if we == "":
            return True
        pal = "".join([i.lower() for i in we if i.isalpha()])
    for i in range (len(pal)):
        if pal[i] != pal [0-(i+1)]:
            return False
    return True
        
        
    file_1.close()
            
is_palindrome ("notPalindrome.txt")
```

Output:
```
False
```

---
### Exercise 2
---
This is a fairly complex function. You will likely want to carefully plan out your design and think about your approach before you begin coding.

There are many different ways to send secret messages using encryption. One popular way is to write a message in which you use an alternative alphabet. It is the same alphabet as the regular English alphabet, but shifted. The alphabet is shifted by 13 letters so that A in the secret code is N, the letter B is O, the letter C is P, and so on. As you keep going, you'll find that M goes to Z, but then what would N go to? N should go to the letter after Z... but there is no letter after Z... so we make the alphabet loop back on itself and say that A comes after Z (kind of like the Ace can come before the 2 or after the King in cards). This code is called ROT-13. ROT-13 has the special property that if you run it again on an encrypted phrase, you get the original back again; for example, encrypting A gives you an N, and encrypting an N gives you A back again. (In general, ROT-n is a shift by n letters; for example, ROT-3, also called Caesar's Cipher because it was invented and used by Julius Caesar a couple millenia ago, encrypts ABCXYZ into DEFABC, while ROT-0 and ROT-26 do not change the message at all. You will only be using ROT-13 in this exercise, though.)

For a more concrete example, the phrase:

    "secret message"
   
will be encrypted via ROT-13 into:

    "frperg zrffntr"

Write a function called **decrypt** that takes one parameter - the name of a file. It should create an output file called **decryptOutput.txt**. The output file should contain the secret message that was decrypted from the file that was given as a parameter. To keep things simple, you may make all of the characters lower case. Be sure to keep all of the  punctuation that was in the original message. In addition to writing it to a file, your function should print the decrypted message to the screen. Your function does not need to return anything.

You should copy the file secretMessage.txt from the ~barbeda/cs128-share/ directory, as you have with other files. Use this file to test with.

**Note:** You will **not** receive full points for writing a function with 26 *if* statements. If you have no idea where to start, think about using the **ord()** function. There is a function that does the opposite of ord, which is **chr()**. Both may come in handy. 

**Hint:** Find the ord of a number, then figure out how to transform the ord of one letter to the ord of a letter that is 13 letters further in the alphabet. Work with mod (%) to do this. 

<!-- grade_id: Part2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def decrypt(name):
    z = (" ")
    file = open(name,"r")
    f = file.read()
    for line in f:
        line.split(",")
        for c in line:
            x=chr((ord(c) - 84) % 26 + 97)
            z += x
    return z
        
        
    file.close()
decrypt("secretMessage.txt")
```

Output:
```
' thisaisatheasecretamessagebadonhtatellaanybodyaaboutaitbee'
```

---
### Exercise 3
---
Write a function called **process_numbers** that opens, reads, analyzes, and closes a file. It should take one parameter, the string containing the file name. The files that this function will analyze are files containing numbers. Use **numbers.dat** as a parameter for testing. (You should copy it to your directory from ~barbeda/cs128-share/numbers.dat) **It should return the following values in this order**:

* the number of numbers that are in the file,
* the total of the numbers,
* the smallest number,
* the largest number, 
* the average of the numbers in that file.

To test the function (this is mandatory), call the function, and then use what it **returns** to **print** something like the following: (The numbers here are not correct; they're just for demonstration.) The function itself should not print anything; the printing should be done by extra code in the cell, outside of the function.

    Count: 5
    Total: 10
    Smallest: 1
    Largest: 10
    Average: 5.5
    
To do this testing you will probably want to start out with something like this:

    count, total, smallest, largest, average = process_numbers("numbers.dat")

<!-- grade_id: Part3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def process_numbers(f):
    anlyz =(["Count:", "Total:", "Smallest:", "Largest:", "Average:"])
    g = []
    file = open(f,"r")
    for line in file:
        lst = line.split(",")
        for i in lst:
            g.append(int(i))
            Count = len(g)
            Total = sum(g)
            Smallest = min(g)
            Largest = max(g)
            Average = Total/Count
    file.close()
    print (anlyz[0],Count)
    print (anlyz[1],Total)
    print (anlyz[2],Smallest)
    print (anlyz[3],Largest)
    print (anlyz[4],Average)
    
process_numbers("numbers.dat")


```

Output:
```
Count: 400
Total: 771710
Smallest: 5
Largest: 3957
Average: 1929.275
```

---
### Exercise 4
---
Design, write and test a cell that reads the file **words.dat** (it's at ~barbeda/cs128-share/words.dat) and determines:

* the shortest line (in number of words)
* the longest line (in number of words)
* the shortest line (by number of characters)
* the longest line (by number of characters)
* the total number of words
* the total number of characters characters

Write a function called **process_words** that opens, reads, analyzes, and closes the file. Your function should take one argument, the name of the file. **The function should return the calculated values in the order above**. It should return multiple values (return them in the order shown above, it will return six values in total). 

To test the function (this is mandatory), call the function, and then use what it **returns** to **print** something like the following: (The numbers here are not correct; they're just for demonstration.)

    Shortest line in words: 5
    Longest line in words: 10
    Shortest line in characters: 40
    Longest line in characters: 80
    Total words: 4010
    Total characters: 39005
    
To do this, you will likely want your test to include something like:

    least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")

<!-- grade_id: Part4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def process_words(name):
    b = []
    anlyz=(["Shortest line in words:", "Longest line in words:", "Shortest line in characters:", "Longest line in characters:", "Total words:", "Total characters:"])
    file = open(name,"r")
    for line in file:
        lst = line.split(",")
        for w in lst:
            b.append(len(w))
        c = len(b)
        t = sum(b)
        s = min(b)
        m = max(b)
        a = t/c
    print (anlyz[0],c)
    print (anlyz[1],t)
    print (anlyz[2],s)
    print (anlyz[3],m)
    print (anlyz[4],a)
    file.close()

process_words("words.dat")            
```

Output:
```
Shortest line in words: 5
Longest line in words: 73
Shortest line in characters: 6
Longest line in characters: 30
Total words: 14.6
```

### Extra Credit
---
Another way to encypt a secret message is to write a message in which the second letter of each word is significant toward the secret message. Write a function called **decode** which takes a coded file with a secret message and prints out the message in an output file named **decodeOutput.txt**. The output file should just contain the letters from the secret message. Don't worry about spaces or punctuation. In addition to printing to the file, your code should display the decoded message to the screen. It does not need to return anything.

**For Example**

If I had a file containing the following:

    aptitude cyclist storybook choice ionic unstable
    
and ran **decode** on it, the output file should be:

    python

and the secret message would be "python"

Use the following file for testing,

    ~barbeda/cs128-share/secretMessage2.txt

<!-- grade_id: Part5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: Test1, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(is_palindrome("224wordPalindrome.txt"), True)
assert_equal(is_palindrome("notPalindrome.txt"), False)
### END TEST
```

<!-- grade_id: Test2, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
decrypt("secretMessage.txt")
### END TEST
```

<!-- grade_id: Test3-1, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

count, total, smallest, largest, average = process_numbers("numbers.dat")

assert_equal(count, 400)
```

<!-- grade_id: Test3-2, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

count, total, smallest, largest, average = process_numbers("numbers.dat")


assert_equal(total, 771710)

### END TEST
```

<!-- grade_id: Test3-3, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

count, total, smallest, largest, average = process_numbers("numbers.dat")


assert_equal(smallest, 5)

### END TEST
```

<!-- grade_id: Test3-4, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

count, total, smallest, largest, average = process_numbers("numbers.dat")


assert_equal(largest, 3957)

### END TEST
```

<!-- grade_id: Test3-5, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

count, total, smallest, largest, average = process_numbers("numbers.dat")


assert_equal(average, 1929.275)
### END TEST
```

<!-- grade_id: Test4-1, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")

assert_equal(least_words, 1)

### END TEST
```

<!-- grade_id: Test-4-2, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")


assert_equal(most_words, 6)

### END TEST
```

<!-- grade_id: Test4-3, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")



assert_equal(least_chars, 6)

### END TEST
```

<!-- grade_id: Test4-4, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")


assert_equal(most_chars, 30)

### END TEST
```

<!-- grade_id: Test4-5, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")


assert_equal(total_words, 17)

### END TEST
```

<!-- grade_id: Test4-6, points: 2, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

least_words, most_words, least_chars, most_chars, total_words, total_chars = process_words("words.dat")


assert_equal(total_chars, 73)
### END TEST
```

<!-- grade_id: ExtraCreditTest, points: 10, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
#Extra Credit
### BEGIN TEST
try:
    decode("secretMessage2.txt")
except (NotImplementedError):
    print("Extra credit was not completed.")
    
### END TEST
```