Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Ihsan Alaeddin"
COLLABORATORS = "Group 1"
```

---

---
# Lab 12: Objects
---
In this lab you will be exploring objects and creating your own objects. There are many implementation details contained below. It may be useful to many of you to print this notebook and then use that as a design space, reference and tracking mechanism as you develop the solutions. 

The utility of working-out the design of your classes ahead of time cannot be emphasized enough. Read through the lab, go through a second time designing the solutions on paper, and then implement and test them on the third pass. 

---
### Exercise 1
---

### Read the instructions carefully!

Write a class called **Card**. (By convention, we capitalize the first letter of class names.) This class will have two attributes (things it knows): suit and rank. The suit and the rank should both be strings. The initializer should take two arguments, the suit and the rank, *in that order*. If you have any questions about how these sorts of playing cards work, please ask us.

The **initializer** for the class should do the following:
* Make sure that the suit of a Card is limited to the following:
        club, spade, diamond, heart (These look like ♣ ♠ ♢ ♡, in that order)
* And that the rank of a Card is limited to the following:
        2, 3, 4, 5, 6, 7, 8, 9, 10, jack, queen, king, ace
* The initialzer should raise a ValueError if either the rank or the suit are not valid.
* Don't worry about joker cards. Your deck should not contain any.

This class will have the following 9 methods:
* One method should be called **get_suit**. It should take no arguments, and return the suit as a String.
* One method should be called **get_rank**. It should take no arguments, and return the rank as a String.
* One method should be called **is_face_card**. It should take no arguments, and return True if the Card is a king, queen, or jack. It should return False otherwise.
* One method should be called **is_red**. It should take no arguments, and return True if the Card is a heart or a diamond. It should return False otherwise. Write a similar function called **is_black** that returns True if the Card is a spade or a club.
* One method should be called **point_value**. It should return an int. That int should be match the rank if the rank is between 2 and 10, inclusive. It should be 1 for aces, 11 for jacks, 12 for queens, and 13 for kings. It should take no arguments. You can use the int() function to convert a string to a numeric value, but only if the string is a series of digits. You can avoid writing 13 if statements by using this.
* One method should be called **is_higher_value_than**. It should take one argument, another Card. It should return True if the Card that is calling the method has a higher point value than the argument, and False otherwise (including if there is a tie.) An exception is that an Ace is higher than anything, even though its point_value is 1. This function should use point_value.
* One method should be called **pretty_string**. It takes no arguments. It returns a 2- or 3-character long string. The first part of the string should be the rank of the card. This is a digit if the rank is a number (10 if it is 10), and A, J, Q, or K if it isn't. The second part of the string should be one of these symbols: ♠ ♣ ♡ ♢. (Those are spades, clubs, hearts, and diamonds, in that order.) Examples of things that might be returned by pretty_string are "10♢", "J♢", "A♡", and "4♠". You may have to copy-and-paste these symbols into your code. They're unicode, so it's okay.
* Write a method called **magic_increment**. It takes no arguments, and doesn't return anything. It transforms the Card into the same Card, but one value higher. An ace will become a 2, a 2 will become a 3, a 10 will become a jack, and a king will become an ace. (Ace is both the highest and lowest.) The suit should remain the same. Consider how you might use your other functions to write this concisely. You can use str() to convert an integer to a string, which may be helpful depending on which other functions you use.

<!-- grade_id: Part1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
class Card():
    
    def __init__(self,suit,rank):
        self.suit = suit
        self.rank = rank

    def get_suit(self):
        a = str(self.suit)
        
    def get_rank(self):
        b = str(self.rank)
    def is_face_card (self):
        if self.rank == "king" or "queen" or "jack":
            return True
        else:
            return False
        
    def is_red(self):
        if self.suit == "heart" or "diamond":
            return True
        else:
            return False
        
    def is_black (self):
        if self.suit == "spade" or "club":
            return True
        else:
            return False
        
    def point_value(self):
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10":
            return int(rank)
        if self.rank == "ace":
            return 1
        if self.rank == "jack":
            return 11
        if self.rank == "queen":
            return 12
        if self.rank == "king":
            return 13
        
    def is_higher_value_than (self,Card):
        if self.rank == "ace":
            return True
        if (self.rank) > (Card):
            return True
        
        else:
            return False
        
    def pretty_string (self):
        
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "spade":
            return (int(rank), "♠" )
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "club":
            return (int(rank), "♣" )
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "heart":
            return (int(rank), "♡" )
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "diamond":
            return (int(rank), "♢" )
        
        if self.rank == "ace" and self.suit == "spade":
            return ("A","♠")
        if self.rank == "ace" and self.suit == "club":
            return ("A","♣")
        if self.rank == "ace" and self.suit == "heart":
            return ("A","♡")
        if self.rank == "ace" and self.suit == "diamond":
            return ("A","♢")
        
        if self.rank == "jack" and self.suit == "spade":
            return ("J","♠")
        if self.rank == "jack" and self.suit == "club":
            return ("J","♣")
        if self.rank == "jack" and self.suit == "heart":
            return ("J","♡")
        if self.rank == "jack" and self.suit == "diamond":
            return ("J","♢")
    
        if self.rank == "queen" and self.suit == "spade":
            return ("Q","♠")
        if self.rank == "queen" and self.suit == "club":
            return ("Q","♣")
        if self.rank == "queen" and self.suit == "heart":
            return ("Q","♡")
        if self.rank == "queen" and self.suit == "diamond":
            return ("Q","♢")
    
        if self.rank == "king" and self.suit == "spade":
            return ("K","♠")
        if self.rank == "king" and self.suit == "club":
            return ("K","♣")
        if self.rank == "king" and self.suit == "heart":
            return ("K","♡")
        if self.rank == "king" and self.suit == "diamond":
            return ("K","♢")
    
    def magic_increment (self):
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "spade":
            return (int(rank)+1, "♠" )
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "club":
            return (int(rank)+1, "♣" )
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "heart":
            return (int(rank)+1, "♡" )
        if self.rank == "2"or"3"or"4"or"5"or"6"or"7"or"8"or"9"or"10" and self.suit == "diamond":
            return (int(rank)+1, "♢" )
        
        if self.rank == "ace" and self.suit == "spade":
            return ("2","♠")
        if self.rank == "ace" and self.suit == "club":
            return ("2","♣")
        if self.rank == "ace" and self.suit == "heart":
            return ("2","♡")
        if self.rank == "ace" and self.suit == "diamond":
            return ("2","♢")
        
        if self.rank == "jack" and self.suit == "spade":
            return ("Q","♠")
        if self.rank == "jack" and self.suit == "club":
            return ("Q","♣")
        if self.rank == "jack" and self.suit == "heart":
            return ("Q","♡")
        if self.rank == "jack" and self.suit == "diamond":
            return ("Q","♢")
    
        if self.rank == "queen" and self.suit == "spade":
            return ("K","♠")
        if self.rank == "queen" and self.suit == "club":
            return ("K","♣")
        if self.rank == "queen" and self.suit == "heart":
            return ("K","♡")
        if self.rank == "queen" and self.suit == "diamond":
            return ("K","♢")
    
        if self.rank == "king" and self.suit == "spade":
            return ("A","♠")
        if self.rank == "king" and self.suit == "club":
            return ("A","♣")
        if self.rank == "king" and self.suit == "heart":
            return ("A","♡")
        if self.rank == "king" and self.suit == "diamond":
            return ("A","♢")
Card()
# Here is some test code you may wish to use, little by little, as you code your class.
card = Card("diamond","ace")
card.get_suit()
card.get_rank()

fiveDiamond = Card("diamond", "5")
aceSpade = Card("spade", "ace")
tenClub = Card("club", "10")
kingHeart = Card("heart", "king")

print(aceSpade.is_higher_value_than(fiveDiamond))
print(aceSpade.is_higher_value_than(tenClub))
print(aceSpade.is_higher_value_than(kingHeart))
print(aceSpade.is_higher_value_than(aceSpade))

print(fiveDiamond.is_higher_value_than(fiveDiamond))
print(fiveDiamond.is_higher_value_than(tenClub))
print(fiveDiamond.is_higher_value_than(kingHeart))
print(fiveDiamond.is_higher_value_than(aceSpade))

print(tenClub.is_higher_value_than(fiveDiamond))
print(tenClub.is_higher_value_than(tenClub))
print(tenClub.is_higher_value_than(kingHeart))
print(tenClub.is_higher_value_than(aceSpade))

print(kingHeart.is_higher_value_than(fiveDiamond))
print(kingHeart.is_higher_value_than(tenClub))
print(kingHeart.is_higher_value_than(kingHeart))
print(kingHeart.is_higher_value_than(aceSpade))
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-102-c24282b691b5> in <module>
    142         if self.rank == "king" and self.suit == "diamond":
    143             return ("A","♢")
--> 144 Card()
    145 # Here is some test code you may wish to use, little by little, as you code your class.
    146 card = Card("diamond","ace")

TypeError: __init__() missing 2 required positional arguments: 'suit' and 'rank'
```

---
### Exercise 2
---
Create a class called **CardDeck**. It should have an *attribute* (a thing it knows) called **deck**, which is a list of **Cards**.

* The initializer should take no arguments. It should define the attribute **deck**, where deck is a list of the 52 playing Cards in order. Each suit should be ordered from Ace to King, and the suits should be in this order: Clubs, Diamonds, Hearts, Spades. This means that the first Card in the deck will be the ace of clubs, and the last will be the king of spades. Your initializer should use for loops or another method to produce the deck variable; you should not write out 52 separate instantiations Card.

* Write a **method** called **shuffle_deck** that shuffles the deck. It should take no arguments. In the `random` library, there is a function called shuffle that you can use. random.shuffle(listName) randomly reorganizes listName. (It does *not* return anything important - it actually changes listName.) In this function, after the deck has been shuffled, *print* a statement communicating that the deck has been shuffled. This function is a mutator, and doesn't need to return anything.

* Write a **method** called **pretty_list** that returns a list of the cards in the deck, in order, using their pretty_string format from the previous exercise. You should call that function; do not rewrite it here. This function will be useful for testing your other functions. This function shouldn't take any arguments.

* Write a **method** called **number_of_cards** that takes no arguments and returns an integer. The integer is the number of Cards in the deck.

* Write a **method** called **cut_deck** that cuts the deck. It should take one argument, a value between 0 and the number of Cards still in the deck (inclusive). We will call that value the cut-point. Cutting the deck should remove all of the Cards above the cut-point and put them at the bottom. (Calling cut_deck with 0 as its argument effectively does nothing.) If the argument is less than 0 or greater than the number of Cards in the deck, your method should raise a valueError.

* Write a **method** called **peek**. It should return the first Card in the deck. (This will be a Card object.) It should take no arguments.

* Write a **method** called **draw**. This function should return the Card at the top of the deck (the beginning of the list) and remove that Card from the deck.  It should take no arguments.

* Write a **method** called **draw_n**. This function should return a list of the n Cards at the top of the deck and remove those Cards from the deck. It should take one argument - n, the number of cards to draw. It should raise a ValueError if more cards are to be drawn than exist in the deck.

* Write a **method** called **contains**. This function should take a Card as an argument, and return True if that Card is still in the deck. This one requires a little bit of clever thinking, because Card objects cannot be compared directly to each other. You'll have to think of a way to compare their attributes. (If you take CS256, you will learn ways to allow them to be compared directly to each other.)

You can use the classes you built in this lab to build a variety of games. If you build something you think is cool, be sure to share it with us.

<!-- grade_id: Part2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import random
class CardDeck ():
    def __init__(self):
        a = -1
        deck1 = []
        deck2 = []
        deck3 = []
        deck4 = []
        suit= [" of Clubs", " of Diamonds", " of Hearts", " of Spades"]
        rank= ["Ace", "Two", "Three", "Four", "Five", "Six", "Seven", "Eight", "Nine", "Ten", "Jack", "Queen", "King"]
        for i in range(13):
            a = a +1
            x = rank[a]+suit[0]
            deck1.append(x)
            
            x = rank[a]+suit[1]
            deck2.append(x)
            
            x = rank[a]+suit[2]
            deck3.append(x)
            
            x = rank[a]+suit[3]
            deck4.append(x)
        
        deck = deck1 + deck2 + deck3 + deck4
        
    def shuffle_deck(self):
        random.shuffle(self.deck)
        print("The Deck has been re-shuffled")
        
    def pretty_list(self):
        return self.deck
    
    def number_of_cards(self):
        return len(self.deck)
    
    def cut_deck(self,number):
        if 0<number<52:
            return self.deck[number:52]
        if number<0 or number>52:
            return ValueError
        
    def peek(self):
        return self.deck[0]
    
    def draw(self):
        return self.deck - self.deck[0]
    
    def draw_n(self,n):
        if 0<number<52:
            return self.deck - self.deck[0:n]
        if number>52:
            return ValueError
    def contains(self,Card):
        if Card == self.deck:
            return True
        else:
            return False

CardDeck ()            
            
    
```

Output:
```
<__main__.CardDeck at 0x7ff1d8142190>
```

<!-- grade_id: initTests, points: 2, grade: true, solution: false -->
```python
#Card
### BEGIN TEST
from nose.tools import assert_raises
# If these tests are failing, your initializer isn't accepting valid inputs
fiveDiamond = Card("diamond", "5")
aceSpade = Card("spade", "ace")
tenClub = Card("club", "10")
kingHeart = Card("heart", "king")
# If these tests are failing, you're not properly raising value errors on bad data passed to the initializer.
assert_raises(ValueError, Card, "diamond", "11")
```

<!-- grade_id: initTest2, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, Card, "daumond", "5")
```

<!-- grade_id: initTest3, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, Card, "diamond", "ten")
```

<!-- grade_id: initTest4, points: 1, grade: true, solution: false -->
```python
assert_raises(ValueError, Card, "diamond", "1")
### END TEST
```

<!-- grade_id: accessorTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(fiveDiamond.get_suit(), "diamond")
assert_equal(aceSpade.get_suit(), "spade")
assert_equal(tenClub.get_suit(), "club")
assert_equal(kingHeart.get_suit(), "heart")

assert_equal(fiveDiamond.get_rank(), "5")
assert_equal(aceSpade.get_rank(), "ace")
assert_equal(tenClub.get_rank(), "10")
assert_equal(kingHeart.get_rank(), "king")

### END TEST
```

<!-- grade_id: characteristicTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST

from nose.tools import assert_true
from nose.tools import assert_false

assert_true(fiveDiamond.is_red())
assert_true(aceSpade.is_black())
assert_true(tenClub.is_black())
assert_true(kingHeart.is_red())

assert_false(fiveDiamond.is_black())
assert_false(aceSpade.is_red())
assert_false(tenClub.is_red())
assert_false(kingHeart.is_black())

assert_false(fiveDiamond.is_face_card())
assert_false(aceSpade.is_face_card())
assert_false(tenClub.is_face_card())
assert_true(kingHeart.is_face_card())

### END TEST
```

<!-- grade_id: pointTests, points: 3, grade: true, solution: false -->
```python
### BEGIN TEST

assert_equal(fiveDiamond.point_value(), 5)
assert_equal(aceSpade.point_value(), 1)
```

<!-- grade_id: pointTests2, points: 1, grade: true, solution: false -->
```python
assert_equal(tenClub.point_value(), 10)
```

<!-- grade_id: pointTests3, points: 1, grade: true, solution: false -->
```python
assert_equal(kingHeart.point_value(), 13)

### END TEST
```

<!-- grade_id: compareTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST

print(aceSpade.rank)
assert_true(aceSpade.is_higher_value_than(fiveDiamond))
assert_true(aceSpade.is_higher_value_than(tenClub))
assert_true(aceSpade.is_higher_value_than(kingHeart))
assert_false(aceSpade.is_higher_value_than(aceSpade))

assert_false(fiveDiamond.is_higher_value_than(fiveDiamond))
assert_false(fiveDiamond.is_higher_value_than(tenClub))
assert_false(fiveDiamond.is_higher_value_than(kingHeart))
assert_false(fiveDiamond.is_higher_value_than(aceSpade))

assert_true(tenClub.is_higher_value_than(fiveDiamond))
assert_false(tenClub.is_higher_value_than(tenClub))
assert_false(tenClub.is_higher_value_than(kingHeart))
assert_false(tenClub.is_higher_value_than(aceSpade))

assert_true(kingHeart.is_higher_value_than(fiveDiamond))
assert_true(kingHeart.is_higher_value_than(tenClub))
assert_false(kingHeart.is_higher_value_than(kingHeart))
assert_false(kingHeart.is_higher_value_than(aceSpade))

### END TEST
```

<!-- grade_id: prettyTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
assert_equal(fiveDiamond.pretty_string(), "5♢")
assert_equal(aceSpade.pretty_string(), "A♠")
assert_equal(tenClub.pretty_string(), "10♣")
assert_equal(kingHeart.pretty_string(), "K♡")

### END TEST
```

<!-- grade_id: magicTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
fiveDiamondMag = Card("diamond", "5")
aceSpadeMag = Card("spade", "ace")
tenClubMag = Card("club", "10")
kingHeartMag = Card("heart", "king")

fiveDiamondMag.magic_increment()
aceSpadeMag.magic_increment()
tenClubMag.magic_increment()
kingHeartMag.magic_increment()

assert_equal(fiveDiamondMag.get_suit(), "diamond")
assert_equal(aceSpadeMag.get_suit(), "spade")
assert_equal(tenClubMag.get_suit(), "club")
assert_equal(kingHeartMag.get_suit(), "heart")

assert_equal(fiveDiamondMag.get_rank(), "6")
assert_equal(aceSpadeMag.get_rank(), "2")
assert_equal(tenClubMag.get_rank(), "jack")
assert_equal(kingHeartMag.get_rank(), "ace")

### END TEST
```

<!-- grade_id: initialDeckTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
# If this test is failing (and the above tests are passing), something is wrong with your contstructor for CardDeck
newDeck = CardDeck()
deck = newDeck.deck
comparableDeck = []
for card in deck:
    comparableDeck.append((card.get_suit(), card.get_rank()))
    
assert_equal(comparableDeck, [('club', 'ace'), ('club', '2'), ('club', '3'), ('club', '4'), ('club', '5'), ('club', '6'), ('club', '7'), ('club', '8'), ('club', '9'), ('club', '10'), ('club', 'jack'), ('club', 'queen'), ('club', 'king'), ('diamond', 'ace'), ('diamond', '2'), ('diamond', '3'), ('diamond', '4'), ('diamond', '5'), ('diamond', '6'), ('diamond', '7'), ('diamond', '8'), ('diamond', '9'), ('diamond', '10'), ('diamond', 'jack'), ('diamond', 'queen'), ('diamond', 'king'), ('heart', 'ace'), ('heart', '2'), ('heart', '3'), ('heart', '4'), ('heart', '5'), ('heart', '6'), ('heart', '7'), ('heart', '8'), ('heart', '9'), ('heart', '10'), ('heart', 'jack'), ('heart', 'queen'), ('heart', 'king'), ('spade', 'ace'), ('spade', '2'), ('spade', '3'), ('spade', '4'), ('spade', '5'), ('spade', '6'), ('spade', '7'), ('spade', '8'), ('spade', '9'), ('spade', '10'), ('spade', 'jack'), ('spade', 'queen'), ('spade', 'king')])


### END TEST
```

<!-- grade_id: PrettyDeckTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
  
assert_equal(newDeck.pretty_list(), ['A♣', '2♣', '3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠'])

### END TEST
```

<!-- grade_id: deckSizeTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
newDeck = CardDeck()
assert_equal(newDeck.number_of_cards(), 52)

### END TEST
```

<!-- grade_id: shuffleDeckTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST

from nose.tools import assert_not_equal

# Technically, this test has an infintisimal chance to fail even if your code is correct, but that's
# probably not what's happening if it fails, especially more than once. (There is a tiny but non-zero)
# chance that a thorougly shuffled deck will end up in the same order as it started; this has never happened in
# history.
deckToShuffle = CardDeck()
deckToShuffle.shuffle_deck()

print("This should be a shuffled deck:")
print(deckToShuffle.pretty_list())
assert_not_equal(newDeck.pretty_list(), deckToShuffle.pretty_list())
assert_equal(deckToShuffle.number_of_cards(), 52)

### END TEST
```

<!-- grade_id: cutDeckTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST
deckToCut = CardDeck()
deckToCut.cut_deck(0)
#print(deckToCut.prettyList())
assert_equal(deckToCut.pretty_list(), ['A♣', '2♣', '3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠'])
deckToCut.cut_deck(1)
assert_equal(deckToCut.pretty_list(), ['2♣', '3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠', 'A♣'])
deckToCut.cut_deck(30)
assert_equal(deckToCut.pretty_list(), ['6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠', 'A♣', '2♣', '3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡'])
deckToCut.cut_deck(52)
assert_equal(deckToCut.pretty_list(), ['6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠', 'A♣', '2♣', '3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡'])
# If these tests fail, it's because you're not raising a ValueError.
assert_raises(ValueError, deckToCut.cut_deck, -1)
assert_raises(ValueError, deckToCut.cut_deck, 53)
### END TEST
```

<!-- grade_id: drawPeekTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST

def prettyHand(cardList):
    prettyCards = []
    for card in cardList:
        prettyCards.append(card.pretty_string())
    return prettyCards

deckToDraw = CardDeck()
firstCard = deckToDraw.draw()
assert_equal(firstCard.pretty_string(),'A♣')
assert_equal(deckToDraw.pretty_list(),['2♣', '3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠'])
secondCard = deckToDraw.draw()
assert_equal(secondCard.pretty_string(),'2♣')
assert_equal(deckToDraw.pretty_list(),['3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠'])
topCard = deckToDraw.peek()
assert_equal(topCard.pretty_string(),'3♣')
assert_equal(deckToDraw.pretty_list(),['3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣', 'J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠'])
nextEight=deckToDraw.draw_n(8)
assert_equal(prettyHand(nextEight),['3♣', '4♣', '5♣', '6♣', '7♣', '8♣', '9♣', '10♣'])
assert_equal(deckToDraw.pretty_list(),['J♣', 'Q♣', 'K♣', 'A♢', '2♢', '3♢', '4♢', '5♢', '6♢', '7♢', '8♢', '9♢', '10♢', 'J♢', 'Q♢', 'K♢', 'A♡', '2♡', '3♡', '4♡', '5♡', '6♡', '7♡', '8♡', '9♡', '10♡', 'J♡', 'Q♡', 'K♡', 'A♠', '2♠', '3♠', '4♠', '5♠', '6♠', '7♠', '8♠', '9♠', '10♠', 'J♠', 'Q♠', 'K♠'])
# If these tests fail, it's because you're not raising a ValueError.
assert_raises(ValueError, deckToDraw.draw_n, 55)
assert_raises(ValueError, deckToDraw.draw_n, 48)
### END TEST
```

<!-- grade_id: containsTests, points: 5, grade: true, solution: false -->
```python
### BEGIN TEST

deckToInspect = CardDeck()
fiveDiamond = Card("diamond", "5")
aceSpade = Card("spade", "ace")
aceClub = Card("club", "ace")

assert_true(deckToInspect.contains(fiveDiamond))
assert_true(deckToInspect.contains(aceSpade))
assert_true(deckToInspect.contains(aceClub))

deckToInspect.draw()

assert_true(deckToInspect.contains(fiveDiamond))
assert_true(deckToInspect.contains(aceSpade))
assert_false(deckToInspect.contains(aceClub))

deckToInspect.draw_n(51)

assert_false(deckToInspect.contains(fiveDiamond))
assert_false(deckToInspect.contains(aceSpade))
assert_false(deckToInspect.contains(aceClub))

### END TEST
```