Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Lab 2: Conversions
**Read instructions carefully.**
---
This lab provides you with practice for using:
1. Functions together,
2. Python packages,
3. Basic strategies for scientific computations.
    
Especially on the more complex exercises, you should attempt to sketch out a design on paper before you code. Be sure to check that your functions are named correctly; it is important that function names match the given specification.

All of the exercises below should use return statements, rather than just printing their outputs. Return statements look like this:

    return output_val
  
Where output_val is the value that the function should return. If you feel confused about the difference between printing and returning, please ask.

Functions defined in one cell can be used in following cells. Once you write a function, you can use it in a later cell, just like you could use a built-in Python functions.

All functions should should take parameters instead of using the input() function. A function that takes two parameters might have a def statement that looks like this:

    def rectangle_area(length, height):



## Exercise 1 - Distances
---
### Exercise 1.1
---
Define a function called **cm_to_in** that takes distance in centimeters **as a parameter** and **returns** the equivalent distance in inches. There are 0.393701 inches in 1 centimeter. It is important that your function **returns** the value that it calculates, rather than printing it.

<!-- grade_id: Part1-1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cm_to_in(cm):
    cm_in_inches = 0.393701*cm 
    return cm_in_inches
cm_to_in(2)
```

Output:
```
0.787402
```

### Exercise 1.2
---
Define a function called **in_to_cm** that takes distance in inches **as a parameter** and **returns** the equivalent distance in centimeters.

<!-- grade_id: Part1-2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def in_to_cm(inc):
    inches_in_cm = 2.54 * inc 
    return inches_in_cm
in_to_cm(0.787402)
```

Output:
```
2.00000108
```

### Exercise 1.3
---
Now, define a function called **in_to_ft** that takes distance in inches **as a parameter** and **returns** the equivalent distance in feet. Remember that there are 12 inches in a foot.

<!-- grade_id: Part1-3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def in_to_ft(ft):
    inches_to_feet = ft/12
    return inches_to_feet 
in_to_ft(2)
```

Output:
```
0.16666666666666666
```

### Exercise 1.4
---
Now, define a function called **ft_to_mi** that takes distance in feet **as a parameter** and **returns** the equivalent distance in miles. There are 5280 ft in 1 mile.

<!-- grade_id: Part1-4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def ft_to_mi(ft):
    feet_to_miles = ft/5280
    return feet_to_miles

ft_to_mi(10560)
```

Output:
```
2.0
```

### Exercise 1.5
---
Define a function called **cm_to_mi** that takes distance in centimeters **as a parameter** and **returns** the equivalent distance in miles. Call the functions that you have written above inside of this function. This function should not do any arithmetic on its own.

<!-- grade_id: Part1-5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def  cm_to_mi(n):
    inches = cm_to_in(n)
    feet = in_to_ft (inches)
    miles = ft_to_mi (feet)
    return (miles)
    

cm_to_mi(1000)    
    
    
```

Output:
```
0.006213715277777778
```

## Exercise 2 - Temperature
---
### Exercise 2.1
---
Kelvin is a temperature scale used by some scientists. One degree in Kelvin is the same as one degree in Celsius, but 0 is in a different place. 0 degrees Kelvin is the coldest anythign can be - absolute zero. Define a function called **K_to_C** that takes temperature in Kelvin **as a parameter** and **returns** the equivalent temperature in degress Celsius. The temperature in Kelvin is 273.15 more than that in degress Celsius.

<!-- grade_id: Part2-1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def K_to_C(k):
    celsius= k-273.15
    return celsius
K_to_C(290)
```

Output:
```
16.850000000000023
```

### Exercise 2.2
---
Define a function called **C_to_F** that takes a temperature in degrees Celsius **as a parameter** and **returns** the equivalent temperature in degrees Fahrenheit. The temperature in degrees Fahrenheit is 32 plus 9/5 of the temperature in degrees Celsius. Check you work by calling your function with a few values to see if the results match your expectations.

<!-- grade_id: Part2-2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def C_to_F(c):
    fahrenheit = c*(9/5) +32
    return fahrenheit
C_to_F(27)
```

Output:
```
80.6
```

### Exercise 2.3
---
Define a function called **K_to_F** that takes temperature in Kelvin **as a parameter** and **returns** the equivalent temperature in Fahrenheit. Use the functions you have defined above - this function should not do any arithmetic on its own.

<!-- grade_id: Part2-3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def K_to_F (f):
    celsius = K_to_C(f)
    fahrenheit = C_to_F(celsius)
    return (fahrenheit)
K_to_F (0)
```

Output:
```
-459.66999999999996
```

## Exercise 3 - Gummi Bears
---
### Exercise 3.1
---
Define a function called **E_to_D** that converts euros to dollars. You can assume that one Euro 
equals 1.11 US dollars. 

<!-- grade_id: Part3-1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def E_to_D(euros):
    dollars=1.11*euros
    return (dollars)
E_to_D(2)
    
```

Output:
```
2.22
```

### Exercise 3.2
---
A 1-pound package of gummy bears is 2.39 dollars. Create a function **D_to_G** that calculates
how many pounds of gummy bears can be bought with a given number of US dollars. This function should take US dollars as a parameter and return a number of pounds of gummy bears.

<!-- grade_id: Part3-2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def D_to_G(dollars):
    pound_of_gummy_bears=dollars/2.39
    return (pound_of_gummy_bears)
D_to_G(4.78)
    
```

Output:
```
2.0
```

### Exercise 3.3
---
Write a function **G_to_C** that takes pounds of gummi bears as a parameter and **returns** the total amount of calories for that amount of gummy bears. A serving of gummy bears is 42 grams 
and has 150 calories. (You will need to look up how many grams are in a pound.)

You will likely want to do this design on paper first in order to keep the units straight.

<!-- grade_id: Part3-3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def G_to_C(pounds):
    grams = 453.592 * pounds
    calories = (150 * grams)/ 42
    return calories

G_to_C(3)   
```

Output:
```
4859.914285714285
```

### Exercise 3.4
---
Stitch the functions you've written so far from exercise 3 together in a function called **E_to_C**. It should return the total amount of gummy bear calories that can be purchased for a given amount of Euros. This function should not do any arithmetic of its own; it should just call other functions. 

**Hint:** This function should take Euros as a parameter

<!-- grade_id: Part3-4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def E_to_C(e):
    dollars = E_to_D(e)
    pounds = D_to_G(dollars)
    calories = G_to_C(pounds)
    return(calories)
    
E_to_C(2)
    
```

Output:
```
1504.743335325762
```

## Exercise 4 - Spheres
---
### Exercise 4.1
---
Write a function named **circle_area** that returns the area of a circle given a radius. 
Rememeber that the area of a circle is pi times the radius squared. 


**Hint:** Python has a way for you to get a relatively precise value for pi without typing in a long decimal number. If you use

    import math

at the top of your code outside the function, you will be able to use a number of useful math functions and variables. You can use

    math.pi
    
as a value of pi. Any time you need the value of pi for a formula, you can use math.pi. The math package has definitions for several variables in it, so when you import the math package, you get access to those variables.

<!-- grade_id: Part4-1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import math
def circle_area(r):
    area=math.pi * r**2
    return area
circle_area(2)
```

Output:
```
12.566370614359172
```

### Exercise 4.2
---
Write a function called **circle_circumference** that takes the radius of a circle as a parameter and returns
the circumference of the circle. The circumference of a circle is pi times the diameter (the diameter is twice the radius). 

<!-- grade_id: Part4-2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def circle_circumference(r):
    circumference= math.pi * r*2
    return circumference
circle_circumference(2)
```

Output:
```
12.566370614359172
```

### Exercise 4.3
---
Define a function called **sphere_volume** that returns the volume of a sphere given the radius.

The formula is:

    V = (4/3) * pi * r^3 (remember to use correct operators)
    

<!-- grade_id: Part4-3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def sphere_volume(r):
    volume = (4/3) * math.pi * r**3
    return volume
    
sphere_volume(2)
```

Output:
```
33.510321638291124
```

### Exercise 4.4
---
Define a function called **sphere_surface_area** that returns the surface area of a sphere given the
radius.

The formula is:

    Surface Area = 4 * pi * (r^2)

<!-- grade_id: Part4-4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def sphere_surface_area(r):
    Surface_Area = 4 * math.pi * (r**2)
    return Surface_Area
sphere_surface_area(2)
```

Output:
```
50.26548245743669
```

### Exercise 4.5
---
Suppose you are making a model of the solar system. You have the spheres for each planet, 
but you need to paint the spheres so that they look like the planets. Utilize the function
**sphere_surface_area** to create a new function called **cost_to_paint** that takes two parameters: the cost of the paint per gram and the radius of the sphere in cm, in that order. The function should return how much it will cost to paint that sphere. Each gram of paint will cover 1.5 square cm. You may assume that the radius is always given in cm.

<!-- grade_id: Part4-5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cost_to_paint(c,r):
    paint = c/1.5
    total_cost = sphere_surface_area(r)*paint
    return total_cost
cost_to_paint(2,2)    
    
```

Output:
```
67.02064327658225
```

### Exercise 4.6
---
Given the radius in meters for a **gold** sphere, create a function **find_mass_of_gold_sphere** which returns the mass of the sphere in kilograms. 
You will need to know the following.

    *The density of gold is 19282 kilograms per cubic meter (kg/m^3)
    
    *Density = Mass / Volume
**Hint**: You already have a function that returns the volume of a sphere - use it. 


<!-- grade_id: Part4-6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def find_mass_of_gold_sphere(r):
    Volume = sphere_volume(r)
    Mass = Volume * 19282
    return Mass
find_mass_of_gold_sphere(2)   
    
```

Output:
```
646146.0218295294
```

### Exercise 4.7
---
Write a function called **cost_of_gold_sphere** that calculates the cost of the gold sphere from exercise 4.6. Since the price of gold fluctuates, the function will take the current price of gold as a parameter along with the radius of the sphere, in that order. Normally, the cost of gold is measured by ounces, so, the cost of gold will always be given in dollars per ounce for this function. You will need to use the function **find_mass_of_gold_sphere** in the new function.


**Hint:** Remember that the function **find_mass_of_gold_sphere** returns the mass in kilograms, you will need to convert kilograms to ounces for this function.


<!-- grade_id: Part4-7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def cost_of_gold_sphere (c,r): 
    total_cost = c*(find_mass_of_gold_sphere(r)*35.274)
    return total_cost
cost_of_gold_sphere (10,1)
```

Output:
```
28490193.46751853
```

## Extra Credit
---
For extra credit, write a function **find_mass_anything**. This function should calculate the 
mass of a sphere with any density. This new function will be very similar to **find_mass_of_gold_sphere**, but it will also take 
the density of the materials the sphere is made from as a parameter. (This means that it will take two parameters: a radius and a density.) The radius will be given in meters, and the density will be given in kilograms per cubic meter.

<!-- grade_id: PartExtraCredit1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def find_mass_anything(r,d):
    Volume = sphere_volume(r)
    Mass = Volume*d
    return Mass
find_mass_anything(637100000,5.51)
    
    
```

Output:
```
5.968470111820102e+27
```

## Extra Credit
---
Write a function called **space_weight** that calculates what an object's weight 
would be if it were on a planet of any size and density. It will take the radius in meters of the planet, 
density of the planet in kilograms per cubic meter, and the mass of the object in kilograms as parameters, in that order. Your function will return the weight in newtons.

To solve this, you will need to calculate the acceleration due to gravity (g) on the planet. 

The formula is:
    g = (G * M)/(r^2)
    
    G is the gravitational constant. If you import the Python library scipy.constants, you can use **scipy.constants.G** as G
    
    M is the mass of the planet you are on and you can find it using **find_mass_anything**
    
    r is the radius of the planet you are on, in meters.


To find out the weight of an object on some planet, multiply the mass of the object times the acceleration due to gravity on the planet. 
    weight = mass * g

If you'd like to know how much Dave weighs on the moon, Dave's mass is 66 kilograms, the density of the moon is 3340 kilograms per cubic meter, and its radius is 1,737,000 meters. Dave's weight on the moon is about 107 Newtons.

**Note**: To give some perspective, on Earth your weight in Newtons equals (your mass in kg) * 9.81. Your weight on the moon should be less that your weight on Earth. 

<!-- grade_id: Part4-8, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import scipy.constants
def space_weight(r,d,m):
    M = find_mass_anything(r,d)
    g = (scipy.constants.G * M)/(r^2)
    Weight=m*g
    return Weight
space_weight(6371,5.51,65)
    
```

Output:
```
4.065470881419153
```

<!-- grade_id: Test1, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(round(cm_to_in(3),2), 1.18)
### END TEST
```

<!-- grade_id: Test1-2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(round(in_to_cm(3),2), 7.62)
### END TEST
```

<!-- grade_id: Test1-3, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(in_to_ft(33), 2.75)
### END TEST
```

<!-- grade_id: Test1-4, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(ft_to_mi(1320), 0.25)
### END TEST
```

<!-- grade_id: Test1-5, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(round(cm_to_mi(93342),2), 0.58)
### END TEST
```

<!-- grade_id: Test2-1, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(round(K_to_C(276),2), 2.85)
### END TEST
```

<!-- grade_id: Test2-3, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(C_to_F(-40), -40)
### END TEST
```

<!-- grade_id: Test2-3-case2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal
### BEGIN TEST
assert_equal(C_to_F(0), 32)
### END TEST
```

<!-- grade_id: Test2-4, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST
assert_equal(round(K_to_F(273.15),2), 32)

### END TEST
```

<!-- grade_id: Test2-4-case2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL

from nose.tools import assert_equal

### BEGIN TEST

assert_equal(round(K_to_F(300),2), 80.33)
### END TEST
```

<!-- grade_id: Test3-1, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(E_to_D(10), 2), 11.1)

###END TEST
```

<!-- grade_id: Test3-1-case2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(E_to_D(91.86), 2), 101.96)
###END TEST
```

<!-- grade_id: Test3-4, points: 1.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(D_to_G(9.01), 2), 3.77)

###END TEST
```

<!-- grade_id: Test3-4-case2, points: 1.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(D_to_G(74.76), 2), 31.28)
###END TEST
```

<!-- grade_id: Test3-4-1, points: 1.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(G_to_C(3.77), 1), 6107.3)

###END TEST
```

<!-- grade_id: Test3-4-1-case2, points: 1.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(G_to_C(31.28), 1), 50672.7)
###END TEST
```

<!-- grade_id: Test3-4-2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(E_to_C(10), 1), 7523.7)

###END TEST
```

<!-- grade_id: Test3-4-2-case2, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(E_to_C(82.76), 1), 62266.3)

###END TEST
```

<!-- grade_id: Test3-4-2-case3, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(E_to_C(1), 1), 752.4)
###END TEST
```

<!-- grade_id: Test4-1, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(circle_area(8), 2), 201.06)

###END TEST
```

<!-- grade_id: Test4-1-case2, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(circle_area(3), 2), 28.27)
###END TEST
```

<!-- grade_id: Test4-2, points: 5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(circle_circumference(5), 2), 31.42)
###END TEST
```

```python

```

<!-- grade_id: Test4-3, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(sphere_volume(5), 2), 523.60)

###END TEST
```

<!-- grade_id: Test4-3-case2, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(sphere_volume(9), 2), 3053.63)
###END TEST
```

<!-- grade_id: Test4-4, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(sphere_surface_area(5), 2), 314.16)

###END TEST
```

<!-- grade_id: Test4-4-case2, points: 0, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(sphere_surface_area(15), 2), 2827.43)
###END TEST
```

<!-- grade_id: Test4-5, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(cost_to_paint(5, .24), 2), 2.41)

###END TEST
```

<!-- grade_id: Test4-5-cse2, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(cost_to_paint(20, .15), 2), 3.77)
###END TEST
```

<!-- grade_id: Test4-6, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(find_mass_of_gold_sphere(.1), 2), 80.77)

###END TEST
```

<!-- grade_id: Test4-6-case2, points: 2.5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal


assert_equal(round(find_mass_of_gold_sphere(2), 2), 646146.02)
###END TEST
```

<!-- grade_id: Test4-7, points: 5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

assert_equal(round(cost_of_gold_sphere(1335.1,.01), 1), 3803.7)
###END TEST
```

<!-- grade_id: Extra1-Test, points: 5, grade: true, solution: false -->
```python
# Ignore this cell

# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

'''3141.5926535897934'''
print(find_mass_anything(5,6))
###END TEST
```

<!-- grade_id: Extra2-Test, points: 5, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from nose.tools import assert_equal

'''107.04594937135553'''
print(space_weight(1737000, 3340, 66))
###END TEST
```